# CODIS: Analyzing STR Profiles and DNA Forensics
### Data Science for Biology

**Notebook developed by:** *Shishi Luo, Sarp Dora Kurtoglu, Kinsey Long*<br>
**Supervised by:** *Steven Brenner*<br>

### Learning Outcomes

In this notebook, you will learn about and practice:
* Intermediate Pandas methods
* Seaborn histograms
* Probability
* Nested for loops
* STRs used for forensics in CODIS

---
#### **Exporting Chats**

When AI tools are used to assist in solving a problem, we will ask for transcripts of any chat sessions as a PDF. Unfortunately, getting Gemini to print a PDF of a chat session while using a CalNet account is not a trivial operation. Here is our suggested solution:

1. Click the `Tools` button at the bottom of the chat box.
2. Select `Canvas`.
3. Copy and paste the following into Gemini:

> Please transcribe our entire conversation word-for-word into a new Canvas document

4. Print to PDF

Alternatively, there are browser-specific extensions that enable chat export. If you choose to use one of these options ***be sure that the extension is one which runs locally, not uploading any data to the cloud***.

#### Use of AI

Throughout this and future labs, we will be explicit about where use of AI is permitted. There will be three possible options, which will be stated up front for each question:

* <font color = #0fd12c>**AI PERMITTED**</font>
* <font color = #d1b70f>**AI REQUIRED**</font>
* <font color = #d14d0f>**NO AI**</font>

#### Submission requirements:

The following are what you are expected to submit:

* Copy of the `.ipynb` notebook file
* PDF version of the notebook file
    * On the DataHub, choose `Export as WebPDF`
* PDF export of Gemini chat logs (if Gemini was used)

---
### Helpful Data Science Resources
Here are some resources you can check out while doing this notebook:

- [Introductory documentation for pandas](https://pandas.pydata.org/docs/user_guide/10min.html)
- [Conditional Probability with Khan Academy](https://www.khanacademy.org/math/ap-statistics/probability-ap/stats-conditional-probability/v/bayes-theorem-visualized)

### Peer Consulting

If you find yourself having trouble with any content in this notebook, Data Peer Consultants are an excellent resource. Click [here](https://dlab.berkeley.edu/training/frontdesk-info) to locate live help.

Peer Consultants are there to answer all data-related questions, whether it be about the content of this notebook, applications of data science in the world, or other data science courses offered at Berkeley.

---

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)

sns.set_theme()

print("All necessary packages have been imported!")
```

## Background
FBI uses a system called CODIS (Combined DNA Index System) to help support and solve criminal cases. The DNA Profile of a suspect is compared against large DNA databases to look for a match. These profiles are determined by DNA short tandem repeat (STR) Analysis with information about allele(s) at the 20 CODIS Core Loci. 

The FBI FAQ about these is [here.](https://www.fbi.gov/resources/dna-fingerprint-act-of-2005-expungement-policy/codis-and-ndis-fact-sheet)<br>
A description of how these are used is [here.](https://www.nature.com/scitable/topicpage/forensics-dna-fingerprinting-and-codis-736)

Until December 31, 2016, there were 13 CODIS Core Loci:

> CSF1PO, FGA, TH01, TPOX, vWA, D3S1358, D5S818, D7S820, D8S1179, D13S317, D16S539, D18S51, D21S11

Later on, 7 more were added, to create a set of 20 CODIS Core Loci:

> CSF1PO, FGA, TH01, TPOX, vWA, D3S1358, D5S818, D7S820, D8S1179, D13S317, D16S539, D18S51, D21S11, D1S1656, D2S441, D2S1338, D10S1248, D12S391, D19S433, D22S1045

In this assignment, we will be using a small set of CODIS data.


## Import the Data

Let's read the raw data in comma separated values format, collected by NIST (National Institute of Standards and Technology) which you can find in the `Data` folder under the name `CODIS_profiles.csv`.  

```python
#Just run this cell
raw_profiles = pd.read_csv("data/CODIS_profiles.csv")
raw_profiles.head()
```

This is an example DNA database used by the FBI. The first row contains headings, including the names of all 20 CODIS Core Loci. Each row after the first one corresponds to an individual and their unique counts for each of the CODIS Core Loci. The alleles here refer to the number of STR repeats in each CODIS Core Loci. Most loci are named as XXXX and XXXX.1. These represent the same locus on the homologous chromosomes. Considering both represents diploid characteristics. For example, the first individual in the dataset has 11 for CSF1PO and 13 for CSF1PO.1. This means that on one chromosome they have 11 repeats at the CSF1PO locus, while on the other chromosome they have 13 repeats at the CSF1PO locus.  Each allele at a locus is named by the number of short tandem repeats; e.g., 6 repeats would be called "allele 6."  (This might seem obvious, but often alleles are named in different ways.) By convention, the smaller number of STR counts is in the XXXX column, and the larger number of STR counts is in the XXXX.1 column.

## Restricting to CODIS original core loci
For this part of the project, we're just going to use the original 13 CODIS loci: TPOX, D3S1358, FGA, D5S818, CSF1PO, D7S820, D8S1179, TH01, VWA, D13S317, D16S539, D18S51, D21S11.

To keep things simple for this part of the lab, we will be considering only one allele from each person, by ignoring the columns with name of the form XXXX.1. We'll also ignore the first two columns which contain ethnic information and the unique identifier.

<font color = #d14d0f>**QUESTION 1**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Create a new DataFrame `CODIS_profiles` which contains data for only the 13 CODIS core loci of each individual. Use the list `CODIS_13_loci`, which contains all 13 STR loci used for DNA fingerprinting as described below, in your code for creating the table. 

```python
# Question 1

CODIS_13_loci = ['TPOX', 'D3S1358', 'FGA', 'D5S818', 
            'CSF1PO', 'D7S820', 'D8S1179', 'TH01', 'vWA',
            'D13S317', 'D16S539', 'D18S51', 'D21S11']

# YOUR CODE HERE
CODIS_profiles = ...

# DO NOT CHANGE CODE BELOW
CODIS_profiles.head()
```

## Exploration and Visualization
The following code plots the distribution of the STR counts for the locus 'TPOX' using the `seaborn` package, abbreviated as `sns`. A histogram can be generated using the figure-level function `sns.displot(data)` (often easier to use for quick plots), or the axes-level function `sns.histplot(data)` (more flexible; more like matplotlib).

You can read the documentation for `sns.displot(data)` [here](https://seaborn.pydata.org/generated/seaborn.displot.html), or for `sns.histplot(data)` [here](https://seaborn.pydata.org/generated/seaborn.histplot.html).


In Data 8, we plotted a histogram using `Table.hist(column_label)`. The primary difference for seaborn is the way the dataframe is input as a parameter, instead of the visualization method being called on the Table object.

An important note is that STR count is a discrete variable, not continuous. Therefore, it is important that we include the `discrete = True` parameter, which centers our x-axis values under the bar.

```python
#Just run this cell
sns.displot(CODIS_profiles, x = "TPOX", discrete=True)
plt.title("Distribution of TPOX STR counts")
plt.xlabel("TPOX STR counts")
plt.ylabel("Number of Individuals")
plt.show()
```

You'll see that the values (the alleles in this case) range from 6 to 11 with a peak at 8. This means that within the population in our DNA dataset, having 8 counts at the TPOX CODIS Core locus is the most common case.

<font color = #d14d0f>**QUESTION 2**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Fill in the code below to create separate bar charts for each CODIS locus (including TPOX), to get a sense of the allelic variation and how concentrated the distributions might be.

Additionally, make sure the plot is legible and interpretable, including the following:

* Make sure each bar chart has appropriate title and axis labels
* Make sure the titles and labels are large enough to be legible
* Tweak the layout so that the subplots fit comfortably on the screen

>**Hint**: Currently, the DataFrame is more wide-form than long-form. How might you reshape the data to be in a form that seaborn prefers? See https://seaborn.pydata.org/tutorial/data_structure.html for guidance.

>**Hint**: [`sns.displot`](https://seaborn.pydata.org/generated/seaborn.displot.html) can perform faceting automatically, and allows you to tweak the layout. Some layout settings are provided in the example code below. Also see the examples in this documentation of how to set titles and labels when faceting.

>**Hint**: Alternatively, you can use [`sns.FacetGrid`](https://seaborn.pydata.org/generated/seaborn.FacetGrid.html) to create subplots and then manually map `sns.histplot` onto that grid. This may afford you more flexibility for fine-tuning how the plots look, at the cost of familarizing yourself with the inner workings of seaborn.


```python
# Question 2

# YOUR CODE HERE
...
g = sns.displot(
    ..., col="Locus",
    col_wrap=..., discrete=..., height=...
)
...

# DO NOT CHANGE CODE BELOW
plt.show()
```

## Upper Bound on the Match Probability
The distribution of alleles for each STR locus in the law enforcement sample is not representative of the entire United States population. When we think about how this data was gathered, only individuals who had crossed paths with the law could have provided sample DNA to build the database. It is fair to assume that there most likely are certain groups that were compelled to provide samples at greater or lesser frequencies contributing to the DNA database. 

However, in an hypothetical situation, if we were to assume that the distribution of alleles for each STR locus in this sample is representative of the population, we can use it to estimate the probability that a suspect's DNA matches the DNA at the crime scene just by chance, when considering only the shorter TPOX STR allele.

Under that hypothetical situation, if we drew individuals at random from the population, it would be reasonable to expect that their alleles for the 'TPOX' STR locus would be distributed approximately as shown below.

```python
#Just run this cell
tpox_counts = CODIS_profiles["TPOX"].value_counts()
tpox_counts
```

The probability that a random individual has a specific allele "x" for a STR locus would be:

$$\frac{\text{number of people with allele } x \text{ in sample}}{\text{number of people in sample}}$$

For example, the most common TPOX allele, allele 8, would be calculated like so:

```python
#Just run this cell
np.round(tpox_counts.max() / tpox_counts.sum(), 3)
```

<font color = #d14d0f>**QUESTION 3a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Suppose that the DNA fingerprint at a crime scene has the most common allele in each of the 13 STR locus in the CODIS set. Assign the probability that a randomly drawn individual would match to the variable `probability_most_common`.

>Note: For this question, and throughout, assume that the probability of having a certain allele at any STR loci is independent of one another. You do not have to follow the skeleton code provided.

```python
# Question 3a

# YOUR CODE HERE
probability_most_common = ...
for locus in CODIS_13_loci:
    frequency_series = ...
    prob_of_locus = ...
    probability_most_common *= ...

# DO NOT CHANGE CODE BELOW
print(f"The probability of a random individual having the most common allele for all 13 STR loci is {probability_least_common:.2e}.")
```

<font color = #d14d0f>**QUESTION 3b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Suppose that the DNA fingerprint at a crime scene has the **least** common allele in the 13 STR locus in the CODIS set. Modify your previous code to assign the variable `probability_least_common` to the probability that a randomly drawn individual would match it.

```python
# Question 3b

# YOUR CODE HERE
probability_least_common = ...
for locus in CODIS_13_loci:
    frequency_series = ...
    prob_of_locus = ...
    probability_least_common *= ...

# DO NOT CHANGE CODE BELOW
print(f"The probability of a random individual having the least common allele for all 13 STR loci is {probability_least_common:.2e}.")
```

<font color = #d14d0f>**QUESTION 3c**:</font> <font color = #d14d0f>**NO AI**</font>

List at least two ways in which the probabilities you calculated above are not fully reflective of reality.

**ANSWER:** Double-click to type your response here.

## Counting the Number of Possible Fingerprints


<font color = #d14d0f>**QUESTION 4a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Suppose this dataset with only one allele per person includes all possible alleles for each STR locus. How many unique DNA fingerprints are possible with the 13 selected STR loci with only one allele per person?

> **Hint**: there is a Series method for counting the number of unique entries.

```python
# Question 4a

# YOUR CODE HERE
unique_STR_combinations = ...
for locus in CODIS_13_loci:
    number_of_alleles = ...
    unique_STR_combinations *= ...

# DO NOT CHANGE CODE BELOW
print(f"The number of unique DNA fingerprints using the 13 STR loci is {f'{unique_STR_combinations:,}'}.")
```

<font color = #d14d0f>**QUESTION 4b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

If we assume that each of these fingerprints are equally likely, what is the probability of a random indivudal having a specific fingerprint, such as one found at a crime scene? Assign your answer to `uniform_prob`. Note that in reality, different fingerprints are not equally likely and often have different probabilities as shown in Question 3.

```python
# Question 4b

# YOUR CODE HERE
uniform_prob = ...

# DO NOT CHANGE CODE BELOW
print(f"The probability of someone having any given fingerprint is {uniform_prob:.3e}.")
```

<font color = #d14d0f>**QUESTION 4c**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

If we assume that each of these fingerprints are equally likely, what is the probability of two random indivudals both sharing a specific fingerprint, such as one found at a crime scene? Assign your answer to `p_two_fingerprints`.

To phrase it in a different way, imagine there is a fingerprint A (the specific given fingerprint). What are the odds that fingerprint B and fingerprint C both match fingerprint A? You can think of the analogy of a given example dice rolling a 3. Now you roll two other dices, what are the chances these two new dices roll the same as the example dice (also roll a 3)? Now, replace dice rolling probability with fingerprint matching probability.

```python
# Question 4c

# YOUR CODE HERE
p_two_fingerprints = ...

# DO NOT CHANGE CODE BELOW
print(f"The probability of two random individuals having a specific fingerprint is {p_two_fingerprints:.3e}")
```

## Comparing fingerprints
For these next two questions, we will keep assuming that each of the fingerprints are equally likely within the population. 

The function below calculates the Euclidean distance between two fingerprints. Euclidean distance is a measure of the straight-line distance between two points in Euclidean space. Euclidean distance is often used as a measure of similarity or dissimilarity between data points. If two points are more dissimilar, they should have a larger distance. There are some limitations of using Euclidean distance - for example, it weights each parameter equally in the distance calculation.

```python
#Just run this cell.
def euclideanDist(fingerprint1, fingerprint2):
    '''given two DNA fingerprints (vectors of the same length) outputs the Euclidean distance'''
    # Note: could use linalg.norm for faster computation, as follows:
    # return np.linalg.norm(fingerprint1 - fingerprint2)
    return np.sqrt(np.sum((fingerprint1 - fingerprint2)**2))
```

<font color = #d14d0f>**QUESTION 5a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Fill in the code below to calculate all pairwise distances between each individual in the dataset.
- Remember that each row, not column, in our DataFrame represents a unique individual.
- Make sure you do not double-count a pairwise distance. Remember that both `euclideanDist(a, b)` and `euclideanDist(b, a)` will calculate the same distance between person `a` and person `b`. Also, you do not want to include redundant distances, like euclideanDist(a, a).

> Note: It may take around a minute for the cell to run.

```python
# Question 5a

# YOUR CODE HERE
distances = ...
for i in range(...):
    for j in range(..., ...):
        distance = euclideanDist(..., ...)
        distances.append(...)
```

<font color = #d14d0f>**QUESTION 5b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Display a bar chart showing the distribution of the pairwise distances. Set the number of bins to 20, and include an appropriate title and axes labels.

```python
#Problem 5b

# YOUR CODE HERE
sns.displot(...)
plt.title(...)
plt.xlabel(...)
plt.ylabel(...)

# DO NOT CHANGE CODE BELOW
plt.show()
```

<font color = #d14d0f>**QUESTION 6a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Adapt your code from Question 5a so that it determines the row indices of the individuals with the shortest Euclidean distance. The row indices of the individuals should be saved in `identity_1` and `identity_2`, and the distance between them should be saved in `shortest_dist`.  

> **Note:** It may take around a minute for the cell to run. In our testing of the lab in 2025, we find that this cell, even when properly coded, may sometimes crash the Kernel. If this happens, please restart and re-run your cells.


```python
#Question 6a

shortest_dist = np.inf #Initializes the shortest_distance as infinity

# YOUR CODE HERE
identity_1 = ...
identity_2 = ...

for i in range(...):
    for j in range(..., ...):
        distance = euclideanDist(..., ...)
        if ...:
            shortest_dist = ...
            identity_1 = ...
            identity_2 = ...

# DO NOT CHANGE CODE BELOW
print(f"""Individuals {identity_1} and {identity_2} have the shortest Euclidean distance of {shortest_dist:.3} units.""")
```

<font color = #d14d0f>**QUESTION 6b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Create the DataFrame `id_comparison` to allow you to manually compare the STR counts of the two individuals determined in Question 6a. The DataFrame indices should be the STR loci, and it should have two columns: `"Person 1"` and `"Person 2"`, which include the STR counts for the individuals. The `"Person 1"` column should contain data from the person with the lower index number.

```python
#Question 6b

# YOUR CODE HERE
id_comparison = pd.DataFrame({"Person 1": ...,
                              "Person 2": ...})

# DO NOT CHANGE CODE BELOW
id_comparison
```

## Diploidy

So far, we only looked at the STR lenght of an allele on one chromosome the shorter alleles for our probability calculations. The original `raw_profiles` DataFrame contains STR alleles for the both chromosomes. The individual's STR count at the same STR locus on their second chromosome is in columns that end with ".1". In this section, we are going to also be considering diploid profiles.  

<font color = #d14d0f>**QUESTION 7a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>


#### Creating a Diploid STR Profile DataFrame

From the original `raw_profiles` DataFrame, you need to create a DataFrame `diploid_profiles` that contains the STR profiles of both chromosomes for each of the original 13 CODIS core loci. 

The `diploid_profiles` DataFrame should have **26 columns**, as each of the original **13 loci** (stored in `CODIS_13_loci`) has two associated values (one for each chromosome). 

##### Step 1: Generate the List of 26 Column Names

First, create a list `CODIS_13_loci_diploid_iterated` that contains the required **26 column names**. This can be done by iterating through `CODIS_13_loci` and appending an additional column name with a `".1"` suffix for each locus.


```python
# Question 7a

CODIS_13_loci_diploid_iterated = []

# YOUR CODE HERE
# Iterate through each locus in CODIS_13_loci to generate the list of 26 column names
for ... in  CODIS_13_loci:
    second_chr_loci = ... + ".1"
    CODIS_13_loci_diploid_iterated.append(...) #First chromosome label
    CODIS_13_loci_diploid_iterated.append(...) #Second chromosome label

# DO NOT CHANGE CODE BELOW
print(CODIS_13_loci_diploid_iterated)
```

<font color = #d14d0f>**QUESTION 7b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>


##### Step 2: Creating the `diploid_profiles` DataFrame

Now, create the DataFrame **`diploid_profiles`** with the needed 26 columns from **`raw_profiles`**, making use of the list **`CODIS_13_loci_diploid_iterated`**.



```python
# Question 7b

# YOUR CODE HERE
# Assign the new columns as part of the Data frame
diploid_profiles = raw_profiles[...]

# DO NOT CHANGE CODE BELOW
diploid_profiles.head()  
```

<font color = #d14d0f>**QUESTION 8**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Create an overlaid bar chart that shows the STR count distribution of both TPOX and TPOX.1. Make sure to set an appropriate binwidth, title, legend, and axes labels.

```python
# Question 8

# YOUR CODE HERE
sns.displot(...)
...

# DO NOT CHANGE CODE BELOW
plt.show()
```

<font color = #d14d0f>**QUESTION 9**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Create a Series `tpox_diploid_counts`, which shows the distribution frequency of each combination of alleles for TPOX. It should have two index columns `"TPOX"` and `"TPOX.1"`, containing the alleles, and its values should represent the number of individuals in `diploid_profiles` with that combination of alleles.

>**Hint**: See the example code for Question 3. There, we used a Series method. There is an equivalent method for DataFrames that will count combinations of columns. How might you ensure that you only use the subset of columns you care about?

>**Note**: You might notice that the output Series shows what appear as blank spaces for "TPOX" in the first column. These are just merged cells that indicate the block of empty rows share the same value. This Series has what is called a [MultiIndex](https://pandas.pydata.org/pandas-docs/stable/user_guide/advanced.html#hierarchical-indexing-multiindex)

```python
# Question 9

# YOUR CODE HERE
tpox_diploid_counts = ...

# DO NOT CHANGE CODE BELOW
tpox_diploid_counts.head(10)
```

<font color = #d14d0f>**QUESTION 10a**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Create a DataFrame `STR_combinations`, which has three columns `"STR locus"`, `"Most Common Allele Combination"` and `"Probability"`. For each of the 13 loci, determine the most commonly observed pair of alleles and the probability (as a proportion) that a random individual would have that pair of alleles. The most common allele combination of a locus should be formatted as "(allele_1, allele_2)" as a string. The probability should be rounded to 2 decimal places.

>**Hints**:
>
> * You may want to consider how you can apply the approach from Question 9.
> * `df.index.values` will extract the indices of a DataFrame as an array (and the same works for Series). If there are multiple layers of indices, they will be output as an array of tuples.
> * You do not have to follow the skeleton code provided.

```python
# Question 10a

# YOUR CODE HERE
locus_tuples = ...
for locus in CODIS_13_loci:
    diploid_counts = ...
    most_common_allele_combination = ...
    most_common_count = ...
    most_common_proportion = ...
    locus_tuples.append(...)

# DO NOT CHANGE CODE BELOW
STR_combinations = pd.DataFrame(
    locus_tuples, columns=["STR locus", "Most Common Allele Combination", "Probability"]
)
STR_combinations
```

<font color = #d14d0f>**QUESTION 10b**:</font> <font color = #0fd12c>**AI PERMITTED**</font>

Calculate the probability that a random individual has the most common allele pair combination for all of the 13 STR loci and assign your result to `probability_most_common_combination`.

```python
# Question 10b

# YOUR CODE HERE
probability_most_common_combination = ...

# DO NOT CHANGE CODE BELOW
print(f"The probability that a random individual has the most common allele pair combination for each of the 13 STR loci is {probability_most_common_combination:.3e}.")
```

<font color = #d14d0f>**QUESTION 10c**:</font> <font color = #d14d0f>**NO AI**</font>

How does your answer in 10a relate to the world population and why does this matter?

**Answer 10c**: YOUR ANSWER HERE!


---

### Congratulations! You have finished the notebook! ##

---

Tests for q1 (points: 1)
```
>>> type(CODIS_profiles) == pd.DataFrame
True
```
```
>>> CODIS_profiles.shape
(1036, 13)
```
```
>>> ' '.join(sorted(CODIS_profiles.columns.to_list()))
'CSF1PO D13S317 D16S539 D18S51 D21S11 D3S1358 D5S818 D7S820 D8S1179 FGA TH01 TPOX vWA'
```


Tests for q10a (points: 6)
```
>>> STR_combinations.shape
(13, 3)
```
```
>>> [type(STR_combinations['STR locus'].iloc[0]), type(STR_combinations['Most Common Allele Combination'].iloc[0])] == [str, str]
True
```
```
>>> len(str(sorted(STR_combinations['Probability'])[0]))
7
```


Tests for q10b (points: 3)
```
>>> 0 < probability_most_common_combination < 1
True
```


Tests for q3a (points: 3)
```
>>> 0 < probability_most_common < 1
True
```


Tests for q3b (points: 2)
```
>>> 0 < probability_least_common < 1
True
```


Tests for q4a (points: 3)
```
>>> 10 ** 6 < unique_STR_combinations < 10 ** 15
True
```


Tests for q4b (points: 1)
```
>>> 0 < uniform_prob < 0.01
True
```


Tests for q4c (points: 2)
```
>>> 0 < p_two_fingerprints <= uniform_prob
True
```


Tests for q5a (points: 5)
```
>>> len(distances)
536130
```
```
>>> np.round(np.mean(distances), 3)
18.6
```


Tests for q6a (points: 4)
```
>>> 0 < shortest_dist < 10
True
```
```
>>> [type(identity_1), type(identity_2)] == [int, int]
True
```


Tests for q6b (points: 2)
```
>>> id_comparison.shape
(13, 2)
```


Tests for q7 (points: 4)
```
>>> diploid_profiles.shape
(1036, 26)
```


Tests for q9 (points: 5)
```
>>> tpox_diploid_counts.shape
(28, 1)
```
```
>>> tpox_diploid_counts.index.to_list()[len(tpox_diploid_counts) // 5]
(8, 10)
```
```
>>> tpox_diploid_counts['count'].iloc[0]
241
```
